<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

<div align=justify>
    <h1 style='text-align:center' ><span style = 'color:#0000FF'>Exercice 1 </span></h1>
    <h2 style='text-align:center' ><i>Epreuve pratique n°2</i></h2>
</div>

<div class = "alert alert-info">
    <b>Q1 : </b>   On considère des mots à trous : ce sont des chaînes de caractères contenant uniquement
des majuscules et des caractères `*`. Par exemple `INFO*MA*IQUE`, `***I***E**` et
`*S*` sont des mots à trous.  
Programmer une fonction correspond qui :

- prend en paramètres deux chaînes de caractères `mot` et `mot_a_trous` où
`mot_a_trous` est un mot à trous comme indiqué ci-dessus, 
- renvoie :
    - `True` si on peut obtenir `mot` en remplaçant convenablement les caractères
`'*'` de `mot_a_trous`.
    - `False` sinon.

Exemple :

```python
>>> correspond('INFORMATIQUE', 'INFO*MA*IQUE')
True
>>> correspond('AUTOMATIQUE', 'INFO*MA*IQUE')
False
```

$\heartsuit$ : On remarquera que si les mots ne sont pas de même longueur ils ne peuvent pas correspondre
</div>

<!-- grade_id: cell-c050cd67976795f4, points: 0, grade: false, solution: true -->
```python
def correspond(mot, mot_a_trous):
    """Vérifie si le mot peut correspondre au mot à trou
    prec : mot et mot_a_trous deux chaines de caractères
    post : Un booléen, True siles motspeuvent correspondre"""
    r = [i for i in range(len(mot)) if mot[i] != mot_a_trous[i]]
    nb = list(mot_a_trous)
    for i in range(0,len(r)):	
        nb[r[i]] = mot[r[i]]
    nb = "".join(nb)
    return True if mot == nb else False
```

<!-- grade_id: cell-b55f83022cb5ddea, points: 1, grade: true, solution: false -->
```python
assert correspond('MATHS','PH*SIQU*') == False
```

<!-- grade_id: cell-210cb2983b48bc3b, points: 2, grade: true, solution: false -->
```python
assert correspond('AUTOMATIQUE', 'INFO*MA*IQUE') == False
```

<!-- grade_id: cell-96a7de615060eb5e, points: 2, grade: true, solution: false -->
```python
assert correspond('INFORMATIQUE', 'INFO*MA*IQUE') == True
assert correspond('NSI','***') == True 
```

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>